---
title: "ixpantia sesion_4: dplyr"
subtitle: "Ejercicios pizarra"
output: 
  learnr::tutorial:
    css: https://storage.googleapis.com/datos_cursos/ixpantia_introR/custom_css.css
runtime: shiny_prerendered
description: >
  Ejercicios de pizarra que se realizan durante la sesión para aprender sobre
  el proceso de resumir datos, específicamente aquellos que involucran el uso
  de dplyr.
---

<style>
@import url('https://fonts.googleapis.com/css2?family=Roboto+Mono&family=Roboto+Slab&family=Source+Sans+Pro:wght@300;400&display=swap');

@import url('https://fonts.googleapis.com/css2?family=Lato&display=swap');
</style> 

```{r setup, include=FALSE}
knitr::opts_chunk$set(echo = FALSE)

library(learnr)
library(dplyr)
library(gradethis)

gapminder <- readr::read_csv('https://raw.githubusercontent.com/OHI-Science/data-science-training/master/data/gapminder.csv')

promedio <- gapminder %>%
    group_by(country) %>%
    summarise(promedio_pib = mean(gdpPercap))

tutorial_options(exercise.checker = gradethis::grade_learnr)
```

```{r, echo=FALSE, out.width="35%", fig.align = "center"}
knitr::include_graphics("https://storage.googleapis.com/datos_cursos/ixpantia_introR/ixpantia_logo.png")  
```

## Datos: **gapminder**

Este conjunto de datos contiene información sobre distintos países tales como
características demográficas o económicas. Incluye 1704 observaciones y
6 variables. 

```{r}
glimpse(gapminder)
```

Descripción de las variables:

| Variable | Descripción |
|:--------:|:-----------:|
| `country` | Nombre del país | 
| `year` | Año de registro |
| `pop`|  Población del país | 
| `continent` | Continente donde se ubica el país |
| `lifeExp` | Esperanza de vida del país |
| `gdpPercap` | PIB per cápita del país |

## Ejercicios 

### Utilice `filter()` para filtrar únicamente los datos de Oceania

En el set de datos **gapminder** hay una columna que se llama `continent` que
muestra el continente al que pertenece cada país. En total tiene 5 categorías.
Filtre de forma que sólo muestre los datos correspondientes a Oceanía usando
`filter()`.

```{r filter, exercise = TRUE}
gapminder %>% 
  filter()
```

```{r filter-hint-1}
"continent == 'Oceania'"
```

```{r filter-check}
grade_result(
  pass_if(~ identical(unlist(.result), 
                      unlist(gapminder %>%
                               filter(continent == 'Oceania'))),
          "¡Muy bien!")
)
```

### Utilice `select()` para seleccionar las columnas `country`, `year` y `pop`

El conjunto de datos **gapminder** tiene en total 6 columnas o variables. Sin
embargo, en ocasiones interesa utilizar una base con solamente algunas
variables.  Por ejemplo podría interesar analizar únicamente la población de
los distintos países en los años registrados. Para lograr esto, seleccione las
columnas `country`, `year` y `pop` utilizando `select()`.

```{r select, exercise = TRUE}
gapminder %>% 
  select()
```

```{r select-check}
grade_result(
  pass_if(~ identical(unlist(.result),
                      unlist(gapminder %>% 
                               select(country, year, pop))),
          "¡Muy bien!")
)
```

### Calcule el promedio del PIB per cápita de cada país usando `summarise()` y `group_by()` 

EL conjunto de datos **gapminder** tiene una variable que se llama `gdpPercap`
que describe el PIB per cápita de los países para distintos años. Ahora utilice
esta variable agrupando por país (`country`) y use `summarise()` para calcular
el promedio de gdpPercap por país. Llame a la nueva variable con los datos del
promedio `promedio_pib`.

```{r summarise, exercise = TRUE}
gapminder %>% 
  group_by() %>% 
  summarise()
```

```{r summarise-hint-1}
"summarise(promedio_gdp = mean(gdpPercap)"
```

```{r summarise-check}
grade_result(
  pass_if(~ identical(unlist(.result),
                      unlist(gapminder %>%
                      group_by(country) %>%
                      summarise(promedio_pib = mean(gdpPercap)))),
          "¡Excelente!")
)
```

### Ordene de menor a mayor los países según `promedio_pib` utilizando `arrange()`

Anteriormente creamos una variable llamada `promedio_pib` que representa el
promedio del PIB per cápita de cada país. Esto quedó guardado en un nuevo
conjunto de datos llamado **promedio**. Ahora, partiendo de este nuevo conjunto
de datos, ordene de de menor a mayor según el PIB promedio de cada país. 

```{r arrange, exercise = TRUE}
promedio %>% 
  arrange()
```

```{r arrange-hint-1}
"Por defecto `arrange()` ordena de menor a mayor"
```

```{r arrange-check}
grade_result(
  pass_if(~ identical(unlist(.result),
                      unlist(promedio %>% 
                      arrange(promedio_pib))), 
          "¡Bien hecho!")
)
```

### Cree una nueva columna que sea la multiplicación de la población por el PIB per cápita usando `mutate()`

En el conjunto de datos **gapminder** hay una columna que se llama `pop` que se
refiere a la población del país y otra columna que se llama `gdpPercap`, la cual
se refiere al PIB per cápita de los países en cada año. Ahora, genere una nueva
variable llamada `multiplicacion` que contenga el valor resultante de la
multiplicación de las dos columnas mencionadas. 

```{r mutate, exercise = TRUE}
gapminder %>% 
  mutate(multiplicacion = )
```

```{r mutate-hint-1}
"* es para multiplicar"
```

```{r mutate-check}
grade_result(
  pass_if(~ identical(unlist(.result),
                      unlist(gapminder %>% 
                      mutate(multiplicacion = pop*gdpPercap))),
          "¡Buenisimo!")
)
```

