---
title: "ixpantia sesion_4: Ejercicios resumir datos"
subtitle: "Ejercicios pizarra"
output: 
  learnr::tutorial:
    css: https://storage.googleapis.com/datos_cursos/ixpantia_introR/custom_css.css
runtime: shiny_prerendered
description: >
  Ejercicios de pizarra que se realizan durante la sesión para aprender sobre
  el proceso de resumir datos tomando las herramientas necesarias vistas
  durante la sesión.
---

<style>
@import url('https://fonts.googleapis.com/css2?family=Roboto+Mono&family=Roboto+Slab&family=Source+Sans+Pro:wght@300;400&display=swap');

@import url('https://fonts.googleapis.com/css2?family=Lato&display=swap');
</style> 

```{r setup, include=FALSE}
knitr::opts_chunk$set(echo = FALSE)

library(learnr)
library(dplyr)
library(tidyr)
library(lubridate)
library(readr)
library(readxl)
library(tibble)
library(gradethis)
tutorial_options(exercise.checker = gradethis::grade_learnr)

directores <- tibble(   
    director =  c("Spielberg",
                  "Scorsese",
                  "Hitchcock",
                  "Tarantino",
                  "Polanski"),    
    nacionalidad = c("EEUU",
                     "EEUU",
                     "RU",
                     "EEUU",
                     "POL"))

peliculas <- tibble(    
    director = c("Spielberg",
		            "Scorsese",
                "Hitchcock",
              	"Hitchcock",
                "Spielberg",
                "Tarantino",
                "Polanski"),    
    pelicula = c("Super 8",
    		        "Taxi Driver",
    		        "Psycho",
    		        "North by Northwest",
    		        "Catch Me If You Can",
    		        "Reservoir Dogs",
    		        "Chinatown"))

premios <- tibble(
  pelicula = c("Catch Me If You Can",
               "Super 8",
               "North by Northwest",
    		       "Chinatown",
    		       "Psycho"),
  premio = c("Oscar",
             "Golden",
             "Oscar",
             "Grammy",
             "Tony"))
```

```{r, echo=FALSE, out.width="35%", fig.align = "center"}
knitr::include_graphics("https://storage.googleapis.com/datos_cursos/ixpantia_introR/ixpantia_logo.png")    
```

## Datos

Para la sesión utilizaremos tres conjuntos de datos distintos:

### **directores**

Contiene información sobre 5 directores de cine. Consta de 2 variables y 5 observaciones.
```{r}
glimpse(directores)
```

Descripción de las columnas:

| Variable       | Descripción                        |
| :----------:   | :--------------------------------: |
| `director`     | Apellido del director              |
| `nacionalidad` | País de donde proviene             |

### **peliculas**

Contiene información sobre 5 películas conocidas. Consta de de 7 observaciones y 2 variables. 
```{r}
glimpse(peliculas)
```

Descripción de las columnas:

| Variable     | Descripción                        |
| :----------: | :--------------------------------: |
| `director`   | Apellido del director              |
| `pelicula`   | Nombre de la película              |

### **premios**

Contiene información sobre los premios de cine ganados por distintas películas. Consta de 5 observaciones y 2 variables. 
```{r}
glimpse(premios)
```

Descripción de las columnas:

| Variable     | Descripción                        |
| :----------: | :--------------------------------: |
| `pelicula`   | Nombre de la película              |
| `premio`     | Nombre del premio que recibió      |

## Ejercicios

### Ejercicio A: Una el conjunto de datos **directores** con **peliculas** y cuente cuántas películas tiene cada director

En el conjunto de datos **directores** tenemos información sobre directores de
cine famosos y su nacionalidad, mientras que en **peliculas** tenemos
información sobre películas conocidas y su director. Ahora una estos dos sets
de forma que podamos calcular el número de películas de cada director. 

```{r A, exercise = TRUE}

```

```{r A-hint-1}
"inner_join"
```

```{r A-check}
grade_result(
  pass_if(~ identical(unlist(.result),
                      unlist(inner_join(directores, peliculas) %>% 
                               group_by(director) %>%
                               tally())),
          "¡Excelente!")
)
```

### Ejercicio B: ¿Cuál es la diferencia entre `right_join` y `left_join`?

Puede utilizar este espacio para hacer ejercicios con `right_join` y
`left_join` antes de contestar la pregunta

```{r B1, exercise = TRUE}

```

```{r B2, echo = FALSE}
question("¿Cuál es la diferencia entre `right_join` y `left_join`?",
  answer("`right_join` no completa la union entre los dos conjuntos de datos"),
  answer("`left_join` une solo las observaciones que están en ambos conjuntos de datos"),
  answer("`right_join` toma solo las observaciones que concuerden con el conjunto de datos de la derecha", correct = TRUE),
  answer("`left_join` toma solo las observaciones que no están en el conjunto de la izquierda"),
  incorrect = "Incorrecto",
  correct = "Correcto")
```

### Ejercicio C: Hacer unión de unicamente las observaciones que concuerdan entre
los conjuntos de datos `peliculas` con `premios`

Utilizando alguno de los *joins* vistos, obtenga un conjunto de datos que tenga
las observaciones que concuerdan en los conjuntos de datos `peliculas` con 
`premios`. Las demás observaciones las debe de excluir

```{r C, exercise = TRUE}

```

```{r C-hint-1}
"inner"
```

```{r C-check}
grade_result(
  pass_if(~ identical(unlist(.result), 
                      unlist(inner_join(peliculas, premios))),
          "¡Excelente!")
)
```

### Ejercicio D: ¿Qué resulta de un `full_join` entre **directores** y **peliculas**?

Al hacer un `full_join` entre los conjuntos de datos **directores** y 
**peliculas**, el conjunto de datos resultante, ¿qué características tendría?

```{r D, echo = FALSE}
question("¿Qué resulta de un `full_join` entre directores y peliculas?",
  answer("Un conjunto de datos con solo las observaciones que coinciden en ambos casos"),
  answer("Un conjunto de datos con las observaciones de directores que coinciden en películas"),
  answer("Un conjunto de datos con las observaciones de películas que coinciden en directores"),
  answer("Un conjunto de datos que contiene la información de ambos sets", correct = TRUE),
  incorrect = "Incorrecto",
  correct = "Correcto")
```

### Ejercicio E: ¿Qué resultaría de un `anti_join` entre **peliculas** y **directores**? 

¿Cuántas observaciones resultarían de aplicar un `anti_join` entre **peliculas** y **directores**?

```{r E, echo = FALSE}
question("¿Qué resultaría de un `anti_join` entre peliculas y directores? ",
  answer("Un set de datos con 12 observaciones"),
  answer("Un set de datos con 5 observaciones"),
  answer("Un set de datos con 7 observaciones"),
  answer("Un set de datos con 0 observaciones", correct = TRUE),
  incorrect = "Incorrecto",
  correct = "Correcto")
```

